<div class="frontmatter text-center">
<h1>Geospatial Data Science</h1>
<h2>Exercise 2: Geospatial Data in Python</h2>
<h3>IT University of Copenhagen, Spring 2023</h3>
<h3>Instructor: Jan Leonard Schelhaas</h3>
</div>

This notebook was adapted from:
* Automating GIS-Processes: https://autogis-site.readthedocs.io/en/latest/lessons/L2/exercise-2.html

## Problem 1: Create Polygon from lists of coordinates

In the first problem **the aim is to create a polygon based on the provided coordinates, and to plot the polygon on as a map.** 

Two lists: `longitudes` and `latitudes` contain the input coordinates for the polygon. The first coordinate pair looks like this: `(29.99671173095703, 63.748023986816406)`.
 
The codeblocks below help you get started with the task. More detailed instructions are given below.

```python
import os
import geopandas as gpd
import matplotlib.pyplot as plt
from shapely.geometry import Polygon

# Longitudes in decimal degrees
longitudes = [29.99671173095703, 31.58196258544922, 27.738052368164062, 26.50013542175293, 26.652359008789062, 25.921663284301758, 22.90027618408203, 23.257217407226562,
           23.335693359375, 22.87444305419922, 23.08465003967285, 22.565473556518555, 21.452774047851562, 21.66388702392578, 21.065969467163086, 21.67659568786621,
           21.496871948242188, 22.339998245239258, 22.288192749023438, 24.539581298828125, 25.444232940673828, 25.303749084472656, 24.669166564941406, 24.689163208007812,
           24.174999237060547, 23.68471908569336, 24.000761032104492, 23.57332992553711, 23.76513671875, 23.430830001831055, 23.6597900390625, 20.580928802490234, 21.320831298828125,
           22.398330688476562, 23.97638702392578, 24.934917449951172, 25.7611083984375, 25.95930290222168, 26.476804733276367, 27.91069221496582, 29.1027774810791, 29.29846954345703,
           28.4355525970459, 28.817358016967773, 28.459857940673828, 30.028610229492188, 29.075136184692383, 30.13492774963379, 29.818885803222656, 29.640830993652344, 30.57735824584961,
           29.99671173095703]

# Latitudes in decimal degrees
latitudes = [63.748023986816406, 62.90789794921875, 60.511383056640625, 60.44499588012695, 60.646385192871094, 60.243743896484375, 59.806800842285156, 59.91944122314453,
           60.02395248413086, 60.14555358886719, 60.3452033996582, 60.211936950683594, 60.56249237060547, 61.54027557373047, 62.59798049926758, 63.02013397216797,
           63.20353698730469, 63.27652359008789, 63.525691986083984, 64.79915618896484, 64.9533920288086, 65.51513671875, 65.65470886230469, 65.89610290527344, 65.79151916503906,
           66.26332092285156, 66.80228424072266, 67.1570053100586, 67.4168701171875, 67.47978210449219, 67.94589233398438, 69.060302734375, 69.32611083984375, 68.71110534667969,
           68.83248901367188, 68.580810546875, 68.98916625976562, 69.68568420410156, 69.9363784790039, 70.08860778808594, 69.70597076416016, 69.48533630371094, 68.90263366699219,
           68.84700012207031, 68.53485107421875, 67.69471740722656, 66.90360260009766, 65.70887756347656, 65.6533203125, 64.92096710205078, 64.22373962402344, 63.748023986816406]

```

**1. Create a Polygon variable `poly` based on the coordinates:**
- Create a list of coordinate pairs based on the provided longitudes and latitudes and store the list in a variable `coordpairs`. A coordinate pair can be either a tuple or a list.
- Create a shapely Polygon using the `coordpairs` -list. Store the polygon in a variable `poly`.

Please use the provided variable names.

<!-- grade_id: create_polygon, points: 0, grade: false, solution: true -->
```python
# Create a list of coordinate pairs
coordpairs = None

# Create a shapely Polygon using the 'coordpairs' -list
poly = None
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

<!-- grade_id: problem_1_polygon_test, points: 0.5, grade: true, solution: false -->
```python
# NON-EDIT CODE CELL FOR TESTING YOUR SOLUTION

# Check the first coordinate pair. Shoud be: (29.99671173095703, 63.748023986816406)
print(coordpairs[0])
```

<!-- grade_id: problem_1_polygon_dtype_test, points: 0.5, grade: true, solution: false -->
```python
# NON-EDIT CODE CELL FOR TESTING YOUR SOLUTION

# Check the data type of your polygon:
print(poly.geom_type)
```

**2. Insert the polygon into a GeoDataFrame called `geo`:**
- Create an empty geodataframe variable `geo`
- Insert our `poly` -polygon into the `geo` GeoDataFrame into a new column named 'geometry'

Please use the provided variable names (replace `None` with your own code).

<!-- grade_id: create_geodataframe, points: 0, grade: false, solution: true -->
```python
# Create an empty GeoDataFrame
geo = None

# Insert the 'poly' -polygon into the 'geo' GeoDataFrame into a new column named 'geometry' 
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

<!-- grade_id: problem_1_gpd_test, points: 0.5, grade: true, solution: false -->
```python
# NON-EDIT CODE CELL FOR TESTING YOUR SOLUTION

# Check the number of rows:
assert len(geo) == 1

assert geo.geom_type.unique()[0] == 'Polygon'

# Check the content of the GeoDataFrame:
geo.head()
```

**3. Plot the polygon.**

Note: you might need to call matplotlib pyplot `show()`-method in order to display the map

<!-- grade_id: plot_map, points: 0, grade: false, solution: true -->
```python
# Plot the polygon. What shape is it :) ?
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

**4. Save the GeoDataFrame into a Geopackage called `'polygon.gpkg'`.**

<!-- grade_id: save, points: 0, grade: false, solution: true -->
```python
# Save the GeoDataFrame into a new Shapefile called 'polygon.shp'.
fp = 'data/polygon.gpkg'

# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

<!-- grade_id: problem_1_plot_and_save_test, points: 2, grade: true, solution: false -->
```python
# NON-EDIT CODE CELL FOR TESTING YOUR SOLUTION

#Check if output file exists
assert os.path.isfile(fp), "Output file does not exist."
```

# Problem 2: Points to map
 
In this problem we continue to learn how to turn latitude and longitude coordinates into geometries.

**Our aim is to plot a map based on a set of longitude and latitude coordinates that are stored in a csv file.** 
The coordinates are in WGS84 decimal degrees (EPSG:4326), and the data is stored in `some_posts.csv` comma separated file in the folder "data". 

First rows of the data look like this:
 
```
lat,lon,timestamp,userid
-24.980792492,31.484633302,2015-07-07 03:02,66487960
-25.499224667,31.508905612,2015-07-07 03:18,65281761
-24.342578456,30.930866066,2015-03-07 03:38,90916112
-24.85461393,31.519718439,2015-10-07 05:04,37959089
```

The data has 81379 rows and consists of locations and times of de-identified social media posts inside Kruger national park in South Africa:

| Column | Description |
|--------|-------------|
| lat | y-coordinate of the post |
| lon | x-coordinate of the post |
| timestamp | Time when the post was uploaded |
| userid | unique id for each user|

*Note: Although the data is based on real social media data, the userids and timestamps have been randomized making it impossible to link the data to the original posts.*

**First:**

- Import the needed modules
- Read the data from `data/some_posts.csv` into a Pandas dataframe called `data`
- Create an empty column called `geometry` where you will store shapely Point objects
- Insert Point objects into the column `geometry` based on the coordinate columns 

<details>
  <summary>Hints</summary>
  
You want to create a Shapely point *on each row*, based on columns `'lon'` and `'lat'`. You can achieve this by using one of these alternative approaches:

- a `for`-loop and `iterrows()` to use the Point constructor on each row
- the [apply method](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.apply.html) to apply the  Point constructor on each row. 
- a `for`-loop and a zipped object containing lon and lat (created using `zip`) and use the Point constructor on each lat, lon coordinate pair.

</details>

<!-- grade_id: Create_points, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

<!-- grade_id: problem_2_input_test, points: 1, grade: true, solution: false -->
```python
# NON-EDIT CODE CELL FOR TESTING YOUR SOLUTION

# Check the result
assert len(data) == 81379
```

<!-- grade_id: problem_2_create_points_test, points: 3, grade: true, solution: false -->
```python
# NON-EDIT CODE CELL FOR TESTING YOUR SOLUTION

# Check the result
assert len(data.loc[data['geometry']== None]) == 0
```

**Next:**
- Convert that DataFrame into a GeoDataFrame

- Set the CRS as WGS84 (i.e. epsg code: 4326)

- Save the data into a geopackage called `Kruger_posts.gpkg`


<!-- grade_id: create_geodataframe, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

<!-- grade_id: problem_2_create_geodataframe_test, points: 1.5, grade: true, solution: false -->
```python
# NON-EDIT CODE CELL FOR TESTING YOUR SOLUTION

# Check that the output file exists
import os
assert os.path.isfile(fp), "output geopackage does not exist"

assert len(geo) == 81379
assert geo.geom_type.unique()[0] == 'Point'
geo.head()
```

**Finally:** 
- Create a simple map of the points using the `plot()`-funtion. 

<!-- grade_id: plot_map, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

## Problem 3: How long have individuals travelled?

In this problem the aim is to calculate the **distance** in meters that the individuals have travelled according to the social media posts (Euclidean distances between points). In this problem, we will need the `userid`-column and the points created in the previous problem. You will need the geopackage `Kruger_posts.gpkg` generated in Problem 2 as input file.

Our goal is to answer these questions based on the input data:
- *What was the shortest distance travelled in meters?*
- *What was the mean distance travelled in meters?*
- *What was the maximum distance travelled in meters?*

Start with importing the required modules and loading the saved data.

<!-- grade_id: read_data, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

1. **Check the CRS of the input data. If this information is missing, set it as EPSG:4326 (WGS84).**
2. **Reproject the data from WGS84-projection into *EPSG:32735*-projection which stands for UTM Zone 35S (UTM zone for South Africa), in order to transform the data into the metric system (don't create a new variable, but update the existing variable `data`)**

<!-- grade_id: reproject, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

<!-- grade_id: problem_3_gdp_test, points: 0.5, grade: true, solution: false -->
```python
# NON-EDIT CODE CELL FOR TESTING YOUR SOLUTION

# Check the data
assert data.crs.to_epsg() == 32735
```

3. **Group the data by userid**

<!-- grade_id: group_data, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

<!-- grade_id: problem_3_grouped_test, points: 0.5, grade: true, solution: false -->
```python
# NON-EDIT CODE CELL FOR TESTING YOUR SOLUTION

#Check the number of groups:
assert len(grouped.groups) == data["userid"].nunique(), "Number of groups should match number of unique users!"
```

**4. Create LineString objects for each user connecting the points from oldest to latest:**

<details>
  <summary>Hints</summary>
  
  **One possible method:**
  - Create an empty DataFrame called `movements`. 
  - Create an empty column "geometry"
  - Use a for-loop where you iterate over the grouped object. For each user's data: 
      - [sort](http://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.sort_values.html) the rows by timestamp 
      - create a LineString object based on the user's points
      - Add the LineString to the geometry column of the `movements` dataframe. You can also add the `userid` in a separate column (or use the userid as index).
  - Convert `movements` into a `GeoDataFrame` (you can replace the DataFrame created in the previous steps with the GeoDataFrame). Remember to set the `geometry` column.
  - Set the CRS of the `movements` GeoDataFrame as `EPSG:32735` 

  OBS! Think about what happens if a visitor only has one entry/location?
  
</details>


<!-- grade_id: create_lines, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

<!-- grade_id: problem_3_movements_tests, points: 3, grade: true, solution: false -->
```python
# NON-EDIT CODE CELL FOR TESTING YOUR SOLUTION

assert movements.crs.to_epsg() == 32735
assert len(movements) == len(data.userid.unique())
#Check the result
print(type(movements))
print(movements.crs)
print(movements['geometry'].head())
```

5. **Finally, calculate the lenghts of the lines into a new column called `distance` in the `movements` GeoDataFrame.**

<!-- grade_id: calculate_distance, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

<!-- grade_id: problem_3_distances_test, points: 2, grade: true, solution: false -->
```python
# NON-EDIT CODE CELL FOR TESTING YOUR SOLUTION

# Check the output - the summed distance should be around 96700
print(f"The visitors in total travelled {movements.dist.sum() / 10000:.0f} km")
```

**You should now be able to answer the following questions:**

- *What was the shortest distance travelled?*
- *What was the mean distance travelled?*
- *What was the maximum distance travelled?*

```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

**Last step is to save the data of into a geopackage called `some_movements.gpkg`**

<!-- grade_id: save_file, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

<!-- grade_id: problem_3_result_tests, points: 0.5, grade: true, solution: false -->
```python
# NON-EDIT CODE CELL FOR TESTING YOUR SOLUTION

import os

#Check if output file exists
assert os.path.isfile(fp), "Output file does not exits."
```